# Data Cleaning and Exploratory Data Analysis

```python
import numpy as
import
```

Tests for q1 (points: [1, 1, 1, 1, 1, 1, 1, 1, 1, 1])
```
>>> q1_1 == 37
True
```
```
>>> q1_2 == 87
True
```
```
>>> q1_3 == 99
True
```
```
>>> q1_4 == [44, 70, 81, 47, 39, 75, 17, 35, 42, 78]
array([ True,  True,  True,  True,  True,  True,  True,  True,  True,
        True])
```
```
>>> import math
>>> math.isclose(q1_5, 52.5, abs_tol=0.01)
True
```
```
>>> q1_6 == 12
True
```
```
>>> import math
>>> math.isclose(q1_7, 59.16, abs_tol=0.01)
True
```
```
>>> q1_8 == [9, 21, 36, 44, 47]
array([ True,  True,  True,  True,  True])
```
```
>>> q1_9 == [9, 12, 9, 14, 9, 0, 1, 0, 3, 13]
array([ True,  True,  True,  True,  True,  True,  True,  True,  True,
        True])
```
```
>>> q1_10 == 3
True
```


Tests for q2 (points: [1, 1, 1, 1, 1, 1, 1, 1])
```
>>> np.allclose(q2_1, [[3.52810469, 0.80031442], [1.95747597, 4.4817864], [3.73511598, -1.95455576], [1.90017684, -0.30271442], [-0.2064377, 0.821197]], atol=1e-05)
True
```
```
>>> np.allclose(q2_2, [3.11188068, 0.95792804, 3.48777285, 0.902668, 0.01065413], atol=1e-05)
True
```
```
>>> np.allclose(q2_3, [0.01555444, 1.85629043, -1.36188064, -0.53595997, 0.02599574], atol=1e-05)
True
```
```
>>> round(q2_4, 5) == round(1.230297949608002, 5)
True
```
```
>>> np.allclose(q2_5, [[1.76405235, 0.40015721], [1.86755799, -0.97727788], [0.95008842, -0.15135721]], atol=1e-05)
True
```
```
>>> round(q2_6, 5) == round(1.8088688373462207, 5)
True
```
```
>>> round(q2_7, 5) == round(2.001256514017904, 5)
True
```
```
>>> np.allclose(q2_8, [0.0, 2.00125651, 1.38131852, 0.98321179, 1.86730039], atol=1e-05)
True
```


Tests for q3 (points: [1, 1, 1, 1, 1, 1, 1])
```
>>> import math
>>> math.isclose(q3_1, -0.07, abs_tol=0.01)
True
```
```
>>> def checkq32(q3_2):
...     tmp = {2017: 0.92, 2018: 0.85, 2019: 0.98}
...     flag = True
...     for i, v in q3_2.items():
...         flag = flag & (round(v, 2) == round(tmp[i], 2))
...     return flag
>>> checkq32(q3_2)
True
```
```
>>> def checkq33(q3_3):
...     tmp = {2017: 1.656, 2018: 1.53, 2019: 1.764}
...     flag = True
...     for i, v in q3_3.items():
...         flag = flag & (round(v, 3) == round(tmp[i], 3))
...     return flag
>>> checkq33(q3_3)
True
```
```
>>> round(q3_4, 5) == round(0.5923333333333335, 5)
True
```
```
>>> q3_5 == 61
True
```
```
>>> q3_6 == 46
True
```
```
>>> def checkq37(q3_7):
...     tmp = {2018: 0.85, 2015: 0.9, 2017: 0.92, 2019: 0.98, 2016: 1.01}
...     flag = True
...     for i, v in q3_7.items():
...         flag = flag & (round(v, 2) == round(tmp[i], 2))
...     return flag
>>> checkq37(q3_7)
True
```


Tests for q4 (points: [1, 1, 1, 1, 1, 1])
```
>>> def checkq41(q4_1):
...     tmp = {'Happiness': 6.9396, 'Log GDP per capita': 10.925769, 'Social support': 0.914219, 'Healthy life expectancy': 68.2995, 'Freedom to make life choices': 0.84262, 'Generosity': 0.149892, 'Perceptions of corruption': 0.699715}
...     flag = True
...     for i, v in q4_1.items():
...         flag = flag & (round(v, 3) == round(tmp[i], 3))
...     return flag
>>> checkq41(q4_1)
True
```
```
>>> def checkq42(q4_2):
...     tmp = {'Happiness': 6.0218, 'Log GDP per capita': 10.340017, 'Social support': 0.93473, 'Healthy life expectancy': 68.604958, 'Freedom to make life choices': 0.877709, 'Generosity': -0.111698, 'Perceptions of corruption': 0.623074}
...     flag = True
...     for i, v in q4_2.items():
...         flag = flag & (round(v, 3) == round(tmp[i], 3))
...     return flag
>>> checkq42(q4_2)
True
```
```
>>> round(q4_3, 5) == round(5.47323986284967, 5)
True
```
```
>>> q4_4 == 53
True
```
```
>>> q4_5 == 57
True
```
```
>>> q4_6 == 5
True
```
