# Data Cleaning and Exploratory Data Analysis

Important note: The 1 million rows dataset on spotify data DOES NOT have genre information. At least when checked most recently... Idk how this was overlooked, but basically since the 1 million songs dataset does not have any genre information, we cannot use that dataset to make analyses on genre or combine with the billboard charts dataset.

```python
import numpy as np
import pandas as pd
from scipy import stats
import seaborn as sns
```

```python
spotify = pd.read_csv("../data/spotify_tracks.csv")
billboard = pd.read_csv("../data/Hot 100.csv")
```

```python
hot_100.head(5)
```

Output:
```
chart_position  chart_date       song           performer  \
0              84  1990-05-05  "B" Girls  Young And Restless   
1              78  1990-05-12  "B" Girls  Young And Restless   
2              68  1990-05-19  "B" Girls  Young And Restless   
3              60  1990-05-26  "B" Girls  Young And Restless   
4              58  1990-06-02  "B" Girls  Young And Restless   

                       song_id  instance  time_on_chart  consecutive_weeks  \
0  "B" GirlsYoung And Restless         1              1                  1   
1  "B" GirlsYoung And Restless         1              2                  2   
2  "B" GirlsYoung And Restless         1              3                  3   
3  "B" GirlsYoung And Restless         1              4                  4   
4  "B" GirlsYoung And Restless         1              5                  5   

   previous_week  peak_position  worst_position chart_debut  \
0            NaN             84              84  1990-05-05   
1           84.0       
```

```python
spotify.head(5)
```

Output:
```
Unnamed: 0                track_id                 artists  \
0           0  5SuOikwiRyPMVoIQDJUgSV             Gen Hoshino   
1           1  4qPNDBW1i3p13qLCt0Ki3A            Ben Woodward   
2           2  1iJBSr7s7jYXzM8EGcbK5b  Ingrid Michaelson;ZAYN   
3           3  6lfxq3CG4xtTiEg7opyCyx            Kina Grannis   
4           4  5vjLSffimiIP26QG5WcN2K        Chord Overstreet   

                                          album_name  \
0                                             Comedy   
1                                   Ghost (Acoustic)   
2                                     To Begin Again   
3  Crazy Rich Asians (Original Motion Picture Sou...   
4                                            Hold On   

                   track_name  popularity  duration_ms  explicit  \
0                      Comedy          73       230666     False   
1            Ghost - Acoustic          55       149610     False   
2              To Begin Again          57       210826     False   
3  
```

```python
# Remove duplicates by track_name and keep the first occurrence
genre = spotify.drop_duplicates(subset="track_name")[["track_name", "track_genre"]]

# Left join to keep all billboard songs and bring in genre info
billboard_genre = pd.merge(billboard, genre, how="left", left_on="song", right_on="track_name")
```

```python
# Distinct songs in billboard
distinct_billboard = billboard.drop_duplicates(subset="song")

# Filter rows where track_genre is not missing (i.e., not NA)
distinct_billboard_genre = billboard_genre[billboard_genre["track_genre"].notna()]

```

```python
# Extract year from chart_date and count songs per year
distinct_billboard_genre["chart_date"] = pd.to_datetime(distinct_billboard_genre["chart_date"])
distinct_billboard_genre["year"] = distinct_billboard_genre["chart_date"].dt.year

songs_per_year = (
    distinct_billboard_genre
    .drop_duplicates(subset="song")
    .groupby("year")
    .size()
    .reset_index(name="count")
    .sort_values("year", ascending=False)
)
songs_per_year

```

Output:
```
C:\Users\natth\AppData\Local\Temp\ipykernel_48960\1742980955.py:2: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  distinct_billboard_genre["chart_date"] = pd.to_datetime(distinct_billboard_genre["chart_date"])
C:\Users\natth\AppData\Local\Temp\ipykernel_48960\1742980955.py:3: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  distinct_billboard_genre["year"] = distinct_billboard_genre["chart_date"].dt.year
    year  count
67  2025     49
66  2024     79
65  2023     60
64  2022    148
63  2021    133
..   ...    
```

```python
# Count distinct songs by genre
songs_per_genre = (
    distinct_billboard_genre
    .drop_duplicates(subset="song")
    ["track_genre"]
    .value_counts()
    .reset_index()
    .rename(columns={"index": "track_genre", "track_genre": "count"})
)
songs_per_genre
```

Output:
```
count  count
0     acoustic    221
1      country    170
2        dance    164
3        blues    158
4     alt-rock    140
..         ...    ...
106      samba      2
107  reggaeton      2
108      indie      2
109      forro      2
110  sertanejo      1

[111 rows x 2 columns]
```

```python
#potentially use this or give up tbh

# spotify_genre_lookup.py
import os
import time
import pandas as pd
from spotipy import Spotify
from spotipy.oauth2 import SpotifyClientCredentials
from tqdm import tqdm

# ---------- CONFIG ----------
# set these in your environment or fill here (env is safer)
# os.environ["SPOTIPY_CLIENT_ID"] = "YOUR_CLIENT_ID"
# os.environ["SPOTIPY_CLIENT_SECRET"] = "YOUR_CLIENT_SECRET"

# initialize Spotipy client (Client Credentials flow)
client_credentials_manager = SpotifyClientCredentials()
sp = Spotify(client_credentials_manager=client_credentials_manager, requests_timeout=10, retries=3)

# files / data frames (adapt paths)
billboard_df = pd.read_csv("data/hot100.csv")    # must contain 'song' and ideally 'artist' columns
# Optional: If you only have song names, still proceed, but include artist when available.

# Create search queries
def make_query(title, artist=None):
    # include quotes to encourage exact title match; artist helps disambiguate
    if pd.isna(artist) or artist == "":
        return f'track:"{title}"'
    return f'track:"{title}" artist:"{artist}"'

# normalize helper
def normalize_string(s):
    if pd.isna(s): return ""
    return " ".join(str(s).lower().strip().split())

# caching helpers
CACHE_FILE = "spotify_genre_cache.parquet"  # persistent cache between runs

if os.path.exists(CACHE_FILE):
    cache = pd.read_parquet(CACHE_FILE)
else:
    cache = pd.DataFrame(columns=["title", "artist", "track_id", "artist_ids", "artist_genres", "match_score", "spotify_found"])

# quick helper to save cache
def save_cache():
    cache.to_parquet(CACHE_FILE, index=False)

# main lookup
results = []

# iterate unique songs (optionally include artist if present)
# reduce to unique pairs for fewer API calls
if "artist" in billboard_df.columns:
    unique_pairs = billboard_df[["song", "artist"]].drop_duplicates().reset_index(drop=True)
else:
    unique_pairs = billboard_df[["song"]].drop_duplicates().reset_index(drop=True)
    unique_pairs["artist"] = ""  # empty artist

for _, row in tqdm(unique_pairs.iterrows(), total=len(unique_pairs)):
    title = str(row["song"])
    artist = str(row["artist"]) if "artist" in row else ""

    qtitle = normalize_string(title)
    qartist = normalize_string(artist)

    # skip if cached
    cached = cache[(cache["title"] == qtitle) & (cache["artist"] == qartist)]
    if not cached.empty:
        results.append(cached.iloc[0].to_dict())
        continue

    query = make_query(title, artist if artist else None)

    # try search
    try:
        r = sp.search(q=query, type="track", limit=5)
    except Exception as e:
        print("Spotify search failed, sleeping and retrying:", e)
        time.sleep(5)
        try:
            r = sp.search(q=query, type="track", limit=5)
        except Exception as e2:
            print("Retry failed:", e2)
            r = {"tracks": {"items": []}}

    items = r.get("tracks", {}).get("items", [])

    if not items:
        # fallback: try searching title only (broader)
        try:
            r2 = sp.search(q=f'track:"{title}"', type="track", limit=10)
            items = r2.get("tracks", {}).get("items", [])
        except Exception:
            items = []

    if not items:
        # nothing found
        entry = {
            "title": qtitle,
            "artist": qartist,
            "track_id": None,
            "artist_ids": None,
            "artist_genres": None,
            "match_score": 0,
            "spotify_found": False
        }
        cache = pd.concat([cache, pd.DataFrame([entry])], ignore_index=True)
        results.append(entry)
        save_cache()
        continue

    # pick the best candidate heuristically (prefer exact artists if available)
    chosen = None
    for it in items:
        # examine track name and artist names
        track_name = normalize_string(it["name"])
        artist_names = [normalize_string(a["name"]) for a in it["artists"]]
        # simple heuristics:
        if track_name == qtitle and (qartist == "" or qartist in artist_names):
            chosen = it
            break
    if chosen is None:
        # fallback to first item
        chosen = items[0]

    track_id = chosen["id"]
    artist_ids = [a["id"] for a in chosen["artists"]]

    # fetch artist genres (batching: you can use sp.artists(list_of_ids))
    try:
        artists_info = sp.artists(artist_ids)
        genres_list = []
        for art in artists_info["artists"]:
            genres_list.extend(art.get("genres", []))
        # unique genres
        genres_unique = list(dict.fromkeys(genres_list))
    except Exception as e:
        print("Failed to get artist info:", e)
        genres_unique = []

    entry = {
        "title": qtitle,
        "artist": qartist,
        "track_id": track_id,
        "artist_ids": artist_ids,
        "artist_genres": genres_unique,
        "match_score": 1,  # heuristically set; could compute similarity
        "spotify_found": True
    }

    cache = pd.concat([cache, pd.DataFrame([entry])], ignore_index=True)
    results.append(entry)

    # persist incrementally to avoid losing progress on long runs
    save_cache()

    # be polite: avoid rate limits (adjust as needed)
    time.sleep(0.12)

# convert results to DataFrame and merge back to Billboard
res_df = pd.DataFrame(results)
# convert genres list to joined string
res_df["genre_join"] = res_df["artist_genres"].apply(lambda g: "; ".join(g) if isinstance(g, list) else None)

# merge back using normalized keys
billboard_df["song_key"] = billboard_df["song"].astype(str).str.lower().str.strip()
billboard_df["artist_key"] = billboard_df["artist"].astype(str).str.lower().str.strip() if "artist" in billboard_df.columns else ""
res_df["title_key"] = res_df["title"]
res_df["artist_key"] = res_df["artist"]

merged = billboard_df.merge(res_df[["title_key","artist_key","genre_join"]], how="left",
                            left_on=["song_key","artist_key"], right_on=["title_key","artist_key"])

merged = merged.drop(columns=["title_key","artist_key"])

merged.to_csv("data/hot100_with_spotify_genres.csv", index=False)
print("Saved merged file with genres (some may be NA).")

```

Output:
```
Cell In[36], line 1
    spotify.loc(contains"B Girls")
                        ^
SyntaxError: invalid syntax
```

Based on this, can we attribute the missing genre-matching to pure-randomness? Is this sample of the billboard 100 chart even though we don't have every song representative? Is it okay to have missing values in this case? Basically, should we use this or should we just scratch the idea of trying to understand how the genres of charting music has changed over time?

Tests for q1 (points: [1, 1, 1, 1, 1, 1, 1, 1, 1, 1])
```
>>> q1_1 == 37
True
```
```
>>> q1_2 == 87
True
```
```
>>> q1_3 == 99
True
```
```
>>> q1_4 == [44, 70, 81, 47, 39, 75, 17, 35, 42, 78]
array([ True,  True,  True,  True,  True,  True,  True,  True,  True,
        True])
```
```
>>> import math
>>> math.isclose(q1_5, 52.5, abs_tol=0.01)
True
```
```
>>> q1_6 == 12
True
```
```
>>> import math
>>> math.isclose(q1_7, 59.16, abs_tol=0.01)
True
```
```
>>> q1_8 == [9, 21, 36, 44, 47]
array([ True,  True,  True,  True,  True])
```
```
>>> q1_9 == [9, 12, 9, 14, 9, 0, 1, 0, 3, 13]
array([ True,  True,  True,  True,  True,  True,  True,  True,  True,
        True])
```
```
>>> q1_10 == 3
True
```


Tests for q2 (points: [1, 1, 1, 1, 1, 1, 1, 1])
```
>>> np.allclose(q2_1, [[3.52810469, 0.80031442], [1.95747597, 4.4817864], [3.73511598, -1.95455576], [1.90017684, -0.30271442], [-0.2064377, 0.821197]], atol=1e-05)
True
```
```
>>> np.allclose(q2_2, [3.11188068, 0.95792804, 3.48777285, 0.902668, 0.01065413], atol=1e-05)
True
```
```
>>> np.allclose(q2_3, [0.01555444, 1.85629043, -1.36188064, -0.53595997, 0.02599574], atol=1e-05)
True
```
```
>>> round(q2_4, 5) == round(1.230297949608002, 5)
True
```
```
>>> np.allclose(q2_5, [[1.76405235, 0.40015721], [1.86755799, -0.97727788], [0.95008842, -0.15135721]], atol=1e-05)
True
```
```
>>> round(q2_6, 5) == round(1.8088688373462207, 5)
True
```
```
>>> round(q2_7, 5) == round(2.001256514017904, 5)
True
```
```
>>> np.allclose(q2_8, [0.0, 2.00125651, 1.38131852, 0.98321179, 1.86730039], atol=1e-05)
True
```


Tests for q3 (points: [1, 1, 1, 1, 1, 1, 1])
```
>>> import math
>>> math.isclose(q3_1, -0.07, abs_tol=0.01)
True
```
```
>>> def checkq32(q3_2):
...     tmp = {2017: 0.92, 2018: 0.85, 2019: 0.98}
...     flag = True
...     for i, v in q3_2.items():
...         flag = flag & (round(v, 2) == round(tmp[i], 2))
...     return flag
>>> checkq32(q3_2)
True
```
```
>>> def checkq33(q3_3):
...     tmp = {2017: 1.656, 2018: 1.53, 2019: 1.764}
...     flag = True
...     for i, v in q3_3.items():
...         flag = flag & (round(v, 3) == round(tmp[i], 3))
...     return flag
>>> checkq33(q3_3)
True
```
```
>>> round(q3_4, 5) == round(0.5923333333333335, 5)
True
```
```
>>> q3_5 == 61
True
```
```
>>> q3_6 == 46
True
```
```
>>> def checkq37(q3_7):
...     tmp = {2018: 0.85, 2015: 0.9, 2017: 0.92, 2019: 0.98, 2016: 1.01}
...     flag = True
...     for i, v in q3_7.items():
...         flag = flag & (round(v, 2) == round(tmp[i], 2))
...     return flag
>>> checkq37(q3_7)
True
```


Tests for q4 (points: [1, 1, 1, 1, 1, 1])
```
>>> def checkq41(q4_1):
...     tmp = {'Happiness': 6.9396, 'Log GDP per capita': 10.925769, 'Social support': 0.914219, 'Healthy life expectancy': 68.2995, 'Freedom to make life choices': 0.84262, 'Generosity': 0.149892, 'Perceptions of corruption': 0.699715}
...     flag = True
...     for i, v in q4_1.items():
...         flag = flag & (round(v, 3) == round(tmp[i], 3))
...     return flag
>>> checkq41(q4_1)
True
```
```
>>> def checkq42(q4_2):
...     tmp = {'Happiness': 6.0218, 'Log GDP per capita': 10.340017, 'Social support': 0.93473, 'Healthy life expectancy': 68.604958, 'Freedom to make life choices': 0.877709, 'Generosity': -0.111698, 'Perceptions of corruption': 0.623074}
...     flag = True
...     for i, v in q4_2.items():
...         flag = flag & (round(v, 3) == round(tmp[i], 3))
...     return flag
>>> checkq42(q4_2)
True
```
```
>>> round(q4_3, 5) == round(5.47323986284967, 5)
True
```
```
>>> q4_4 == 53
True
```
```
>>> q4_5 == 57
True
```
```
>>> q4_6 == 5
True
```
